Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your`name` and `student ID` below:

```python
NAME = ""
STUDENT_ID = ""
```

---

## Exercises - Introduction to Python Programming

---
## No. 1

Complete the function below that accepts 5 numbers as arguments and returns the average of those numbers.

<!-- grade_id: calc_avg, points: 0, grade: false, solution: true -->
```python
def calc_avg(a,b,c,d,e):
    """Compute the average of five (5) numbers and return the result."""
    # YOUR CODE HERE
    
```

Your function should return `3.0` for `calc_avg(1,2,3,4,5)`. Check that it does:

```python
calc_avg(1,2,3,4,5)
```

<!-- grade_id: correct_calc_avg, points: 2, grade: true, solution: false -->
```python
"""Check that calc_avg returns the correct output for several inputs"""
assert calc_avg(1,2,3,4,5) == 3.0
assert calc_avg(10,21,30,42,50) == 30.6
assert calc_avg(10.34,19,30,80,120) == 51.86800000000001
```

---

## No. 2 

Complete the function below that takes accepts a single integer and displays whether the integer is odd or even.

<!-- grade_id: is_odd_even, points: 0, grade: false, solution: true -->
```python
def is_odd_even(n):
    """ Check wether the input n is odd or even."""
    # YOUR CODE HERE
    
```

The `is_odd_even` function should return 'even' when the input is 10. Verify that this is the answer you get:

```python
is_odd_even(10)
```

<!-- grade_id: correct_is_odd_even, points: 2, grade: true, solution: false -->
```python
"""Check that is_odd_even returns the correct answer for various inputs."""
assert is_odd_even(10) == 'even'
assert is_odd_even(15) == 'odd'
assert is_odd_even(89) == 'odd'
assert is_odd_even(12004) == 'even'
```

---
## No. 3

Complete the function below that accepts two arguments, $\texttt{a}$ and $\texttt{b}$, and returns the remainder of $\texttt{a/b}$. 

<!-- grade_id: remainder_ab, points: 0, grade: false, solution: true -->
```python
def remainder_ab(a,b):
    """ Returns the remainder of a/b."""
    # YOUR CODE HERE
    
```

The `remainder_ab` function should return 1 when the value of a=10 and b=3. Verify that this is the answer you get:

```python
remainder_ab(10,3)
```

<!-- grade_id: correct_remainder_ab, points: 2, grade: true, solution: false -->
```python
"""Check that remainder_ab returns the correct answer for various inputs."""
assert remainder_ab(10,3) == 1
assert remainder_ab(920,5) == 0
assert remainder_ab(12345,7) == 4
assert remainder_ab(12943,13) == 8
```

---
## No. 4 

Complete the function below that accepts a single integer $N$, as an argument, and returns the number of factors of $N$.

<!-- grade_id: check_fact, points: 0, grade: false, solution: true -->
```python
def check_fact(n):
    """ Returns the number of factors of n."""
    # YOUR CODE HERE
    
```

The factor of a number is a number that divides the given number completely without any remainder. The factors of a number can be positive or negative. For If the function receives 8 as an argument, it should return 4.example, let us find the factors of 8. Since 8 is divisible by 1, 2, 4, and 8, we can list the positive factors of 8 as, 1, 2, 4, and 8. If the `check_fact` function receives 8 as an argument, it should return 4.

```python
check_fact(8)
```

<!-- grade_id: correct_check_fact, points: 4, grade: true, solution: false -->
```python
"""Check that check_fact returns the correct answer for various inputs."""
assert check_fact(8) == 4
assert check_fact(99) == 6
assert check_fact(1937) == 4
assert check_fact(25) == 3
```